---
  title: "Template Chapter 1"
  description: "This is a template chapter."
---

## Ex 1.1

```yaml
type: NormalExercise 
lang: sql
xp: 100 
skills: 1
key: b28b7c216c   
```


Do some data science.


---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: ff2923c4b7   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`

```{python}

```

